Causal evidence for a domain-specific role of left superior frontal sulcus in human perceptual decision making

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

In this study, participants completed two different tasks. A perceptual choice task in which they compared the sizes of pairs of items and a value-different task in which they identified the higher value option among pairs of items with the two tasks involving the same stimuli. Based on previous fMRI research, the authors sought to determine whether the superior frontal sulcus (SFS) is involved in both perceptual and value-based decisions or just one or the other. Initial fMRI analyses were devised to isolate brain regions that were activated for both types of choices and also regions that were unique to each. Transcranial magnetic stimulation was applied to the SFS in between fMRI sessions and it was found to lead to a significant decrease in accuracy and RT on the perceptual choice task but only a decrease in RT on the value-different task. Hierarchical drift diffusion modelling of the data indicated that the TMS had led to a lowering of decision boundaries in the perceptual task and a lower of nondecision times on the value-based task. Additional analyses show that SFS covaries with model derived estimates of cumulative evidence, that this relationship is weakened by TMS.

Strengths:

The paper has many strengths, including the rigorous multi-pronged approach of causal manipulation, fMRI and computational modelling, which offers a fresh perspective on the neural drivers of decision making. Some additional strengths include the careful paradigm design, which ensured that the two types of tasks were matched for their perceptual content while orthogonalizing trial-to-trial variations in choice difficulty. The paper also lays out a number of specific hypotheses at the outset regarding the behavioural outcomes that are tied to decision model parameters and well justified.

We thank the reviewer for their thoughtful summary of the study and for highlighting these strengths. We are pleased that the multi-pronged approach combining causal manipulation, fMRI, and hierarchical drift–diffusion modelling, as well as the careful matching of perceptual content across the two tasks, came across clearly. We also appreciate the reviewer’s positive remarks on the specificity of our a priori hypotheses and their links to decision-model parameters. In revising the manuscript, we have aimed to further streamline the presentation of these hypotheses and to more explicitly connect the behavioural predictions, model parameters, and neural readouts throughout the Results and Discussion sections.

Weaknesses:

In my previous comments (1.3.1 and 1.3.2) I noted that key results could be potentially explained by cTBS leading to faster perceptual decision making in both the perceptual and value-based tasks. The authors responded that if this were the case then we would expect either a reduction in NDT in both tasks or a reduction in decision boundaries in both tasks (whereas they observed a lowering of boundaries in the perceptual task and a shortening of NDT in the value task). I disagree with this statement. First, it is important to note that the perceptual decision that must be completed before the value-based choice process can even be initiated (i.e. the identification of the two stimuli) is no less trivial than that involved in the perceptual choice task (comparison of stimulus size). Given that the perceptual choice must be completed before the value comparison can begin, it would be expected that the model would capture any variations in RT due to the perceptual choice in the NDT parameter and not as the authors suggest in the bound or drift rate parameters since they are designed to account for the strength and final quantity of value evidence specifically. If, in fact, cTBS causes a general lowering of decision boundaries for perceptual decisions (and hence speeding of RTs) then it would be predicted that this would manifest as a short NDT in the value task model, which is what the authors see.

We thank the reviewer for raising these points and for the helpful clarification. We agree that, in principle, the architecture of the value-based task can be conceived as involving an upstream perceptual process that must be completed, to some degree, before value comparison can proceed. Under such a multistage framework, it is indeed possible that cTBS-induced changes in a perceptual decision stage could manifest as a reduction in boundary separation in the pure perceptual task, while the same perturbation appears as a shortening of non-decision time (NDT) when fitting a single-stage DDM to the value task. In this sense, our earlier statement that a “general speeding effect” would necessarily produce identical parameter changes (either NDT or boundaries) in both tasks was too strong, and we are grateful to the reviewer for pointing this out.

At the same time, this alternative explanation remains fully compatible with our central claim that the left SFS plays a perceptual rather than value-based role. We agree with the reviewer that there must be a stimulus-related circuit (in visual and parietal regions) that encodes the physical attributes of the options, and that this upstream processing can influence both tasks. However, a large body of work suggests that left SFS is not part of this primary identification circuitry, but rather contributes specifically to the accumulation and comparison of sensory evidence (e.g., Heekeren et al., 2004, 2006), downstream from areas such as FFA, PPA, or MT/V5 that encode stimulus identity. In other words, stimulus identification (forming a representation of “what is where”) is anatomically and functionally distinct from the accumulation of evidence toward a perceptual decision. Within this framework, the reviewer’s proposal that cTBS speeds “perceptual decisions” across tasks can be understood as targeting precisely the evidence-accumulation stage we ascribe to SFS, with the value-comparison stage proper likely implemented in other regions (e.g., vmPFC and connected valuation circuitry).

We therefore do not rely solely on the dissociation between boundary changes in the perceptual task and NDT changes in the value task as decisive evidence against a “general speeding” account. Instead, our interpretation is based on the convergence of behavioural, model-based, and neural results. First, in the perceptual task, cTBS to left SFS leads to a selective reduction in decision boundary and a concomitant change in trialwise BOLD activity within the stimulated region that covaries with perceptual choice behaviour and with the latent decision variable inferred from the HDDM. Second, in the value task, cTBS does not affect value sensitivity or accuracy, nor does it alter value-related drift or boundary parameters; the only robust HDDM effect is a modest shortening of NDT. Third, critically, left SFS BOLD activity is modulated by perceptual evidence and by cTBS in the perceptual task, but we observe no evidence that SFS activity encodes value evidence or shows value-related cTBS neuronal effects in the value task.

Taken together, these findings indicate that the left SFS serves a causal role in the accumulation of perceptual evidence and in the setting of the choice criterion for perceptual decisions. The reviewer’s suggestion that cTBS may induce a general speeding of perceptual processes that also influences the value task is compatible with this conclusion, in the sense that any contribution of SFS to the value task is best understood as acting via a perceptual component that is upstream of value comparison, rather than via the value accumulation process itself. We have clarified this point in the Discussion of the revised manuscript and now explicitly acknowledge that our DDM dissociation alone does not exclude a general perceptual speeding account, but that the combination of task-specific neural effects in SFS, preserved value-based choice behaviour, and the absence of value-related BOLD changes in SFS strongly support a primarily perceptual role for this region.

Reviewer #2 (Public review):

Summary:

The authors set out to test whether a TMS-induced reduction in excitability of the left Superior Frontal Sulcus influenced evidence integration in perceptual and value-based decisions. They directly compared behaviour-including fits to a computational decision process model---and fMRI pre and post TMS in one of each type of decision-making task. Their goal was to test domain-specific theories of the prefrontal cortex by examining whether the proposed role of the SFS in evidence integration was selective for perceptual but not value-based evidence.

Strengths:

The paper presents multiple credible sources of evidence for the role of the left SFS in perceptual decision making, finding similar mechanisms to prior literature and a nuanced discussion of where they diverge from prior findings. The value-based and perceptual decision-making tasks were carefully matched in terms of stimulus display and motor response, making their comparison credible.

We thank the reviewer for their clear summary of our aims and approach, and for highlighting these strengths. We are pleased that the convergence between causal TMS, fMRI, and hierarchical modelling comes across as providing credible evidence for the role of left SFS in perceptual decision-making, and that our attempt to link these results to the existing literature is seen as appropriately nuanced. We also appreciate the reviewer’s positive assessment of the task design, in particular the close matching of perceptual content and motor output across perceptual and value-based decisions, which was central to our goal of testing domain-specific theories of prefrontal function. In revising the manuscript, we have further clarified these design choices and their rationale, and we have streamlined the exposition of how the hypotheses, model parameters, and neural readouts are connected across the two decision domains.

Weaknesses:

I was confused about the model specification in terms of the relationship between evidence level and drift rate. While the methods (and e.g. supplementary figure 3) specify a linear relationship between evidence level and drift rate, suggesting, unless I misunderstood, that only a single drift rate parameter (kappa) is fit. However, the drift rate parameter estimates in the supplementary tables (and response to reviewers) do not scale linearly with evidence level.

We thank the reviewer for raising this point and appreciate the opportunity to clarify the model specification. In our hierarchical DDM, we did not fit separate, free drift parameters for each evidence level. As shown in Supplementary Fig. 3, the drift on each trial is specified as

where 𝐸𝑐,𝑠,𝑖 the trial-wise evidence (difference in size or value) and κ𝑐,𝑠 is a single drift-scaling parameter per condition and session. Thus, the linear dependence of drift on evidence is implemented at the trial level via 𝜅; we do not estimate independent 𝛿 parameters for each evidence level.

In Supplementary Tables 8 and 9 we report, for descriptive purposes, the posterior means of 𝛿 conditional on each evidence bin (levels 1–4), alongside the corresponding decision boundary and nondecision time summaries. These values are therefore derived quantities that reflect the combination of (i) the single κ𝑐,𝑠 parameter, (ii) the empirical distribution of continuous evidence values 𝐸 within each bin, and (iii) hierarchical pooling across subjects and sessions. Consequently, they are expected to increase monotonically with evidence level—as they do in our data—but not to lie exactly on a straight line in the discrete level index, because the underlying evidence bins are not equally spaced in physical units and because of between-subject variability and posterior uncertainty.

We will revise the text and table captions to make clear that the evidence-level entries are descriptive summaries of 𝛿 implied by the 𝜅×𝐸 formulation, rather than independently estimated drift parameters, in order to avoid this confusion.

-The fit quality for the value-based decision task is not as good as that for the PDM, and this would be worth commenting on in the paper.

We agree that the HDDM fit for the value-based task is somewhat weaker than for the perceptual task. This is reflected in the somewhat higher DIC values for VDM compared with PDM and in slightly broader posterior-predictive distributions (Supplementary Tables 8–11 and Supplementary Figs. 11–16). We believe this difference primarily reflects the greater intrinsic variability of subjective value-based choices (e.g. trial-to-trial fluctuations in preferences, satiety, or attention), coupled with our decision to use the same relatively simple DDM architecture for both tasks to allow a principled cross-task comparison. Importantly, posterior-predictive checks show that, for VDM as well, the model adequately reproduces both accuracy and full RT distributions at the group and subject level (Supplementary Figs. 11–16), indicating that the fit quality is sufficient for our purposes. In the revised manuscript we now explicitly note that the model captures PDM behaviour more tightly than VDM and that this may reduce sensitivity to very small cTBS effects on value-based decision parameters, even though no systematic effects are evident in our data. Crucially, our central conclusion—that left SFS plays a domain-specific role in setting the decision boundary for perceptual evidence—relies on the robust behavioural, computational, and neural effects observed in PDM and does not depend on assuming a perfect model fit for VDM.

- Supplementary Figure 3 specifies the distribution for kappa hyper-parameter twice.

We thank the reviewer for spotting this typo. We have revised Supplementary Figure 3 legend.

Changes from the earlier version to the revised version

Abstract

Humans and animals can flexibly choose their actions based on different information, ranging from objective states of the environment (e.g., apples are bigger than cherries) to subjective preferences (e.g., cherries are tastier than apples). Whether the brain instantiates these different choices by recruiting either specialised or shared neural circuitry remains debated. Specifically, domain-general accounts of prefrontal cortex (PFC) function propose that prefrontal areas flexibly process either perceptual or value-based evidence depending on what is required for the present choice, whereas domain-specific theories posit that PFC sub-areas, such as the left superior frontal sulcus (SFS), selectively integrate evidence relevant for perceptual decisions. Here we comprehensively test the functional role of the left SFS for choices based on perceptual and value-based evidence, by combining fMRI with a behavioural paradigm, computational modelling, and transcranial magnetic stimulation. Confirming predictions by a sequential sampling model, we show that TMS-induced excitability reduction of the left SFS selectively changes the processing of decision-relevant perceptual information and associated neural processes. In contrast, value-based decision making and associated neural processes remain unaffected. This specificity of SFS function is evident at all levels of analysis (behavioural, computational, and neural, including functional connectivity), demonstrating that the left SFS causally contributes to evidence integration for perceptual but not value-based decisions.

Introduction

Humans and animals alike perform a mélange of goal-directed decisions that require the accumulation of different types of information. If the goal, for example, is to accurately determine whether an apple is bigger than a cherry (perceptual choice), the decision-maker accumulates size information of each fruit; or, the decision-maker may draw out information from personal taste profiles if the goal is to determine whether consuming a cherry over an apple maximises their subjective preferences (value-based choice). Previous studies have shown that different brain circuitries are recruited to accumulate evidence that would instantiate such distinct goal-directed decisions (Summerfield and Tsetsos, 2012; Polanía et al., 2014, 2015; Grueschow et al., 2015); thus, it remains debated to what degree certain decision-making processes share neural circuitry or whether these processes operate under specialised systems. However, prior studies were largely correlational (Heekeren et al. 2004, 2008; Polanía et al., 2014; Grueschow et al., 2015), and most causal studies were only limited to one type of choice (Philiastides et al., 2011; Rahnev et al., 2016) and performed in animals (Ding and Gold, 2012b; Erlich et al., 2015; Hanks et al., 2015). Animal studies provide critical causal insights, yet direct translation to humans can be limited by species-specific anatomy and potential non-homologies (e.g., human SFS versus frontal orienting fields in rodents). Therefore, establishing causal contributions in the human brain remains essential.

In the absence of a comparison choice task, it is impossible to ascertain whether neural circuitry is domain-specific to a particular process, or domain-general that it may be involved across many types of choices. Very few studies (Polanía et al., 2014, 2015; Grueschow et al., 2015) have carefully matched perceptual and value-based decisions in terms of evidence strength, stimulus display, and response modality, and compare them through the lens of a common sequential-sampling framework of evidence accumulation (Dutilh and Rieskamp, 2016; Gold and Shadlen, 2007; Krajbich, 2019), which has long been applied to both perceptual (Ratcliff and Mckoon, 1988) as well as value-based (Busemeyer and Townsend, 1993; Usher and McClelland 2001) decisions. Such studies were able to identify common and specialised circuitries and mechanisms associated to perceptual or value-based decisions or both (Polanía et al., 2014; Grueschow et al., 2015).

But unless causality is established, it is even more difficult to attribute the circuitry’s role in evidence accumulation for one or several choice domains, or whether its involvement is peripheral and merely functionally supporting a larger system. Given task complexity, such studies of observing causal neural effects in healthy human populations using non-invasive brain stimulation are incredibly sparse. One previous study has, at least, shown that causally de-synchronising frontoparietal connectivity specifically increased choice variability during value-based choice, but had no effect on perceptual decisions (Polanía et al., 2015); thus, establishing the causal role of the frontoparietal network during value-based choice. But while indeed causal, the study was limited, relative to the standards of evidence in animal studies (Erlich et al., 2015; Hanks et al., 2015; Piet et al., 2017), since its results, as in many causal stimulation studies in humans (Philiastides et al., 2011; Rahnev et al., 2016), showed behavioural, but no neural effects. Furthermore, we only have evidence of a single dissociation that shows a causal stimulation effect specific to value-based, and not perceptual choice. What candidate region would show a causal effect that is specific to perceptual, not value-based decisions in a way that would demonstrate a double dissociation?

Seminal human imaging studies have repeatedly implicated the superior frontal sulcus (SFS), a posterior portion of the dorsolateral prefrontal cortex (dlPFC), during perceptual decision-making (Heekeren et al., 2004, 2008; Mulder et al., 2014). While these studies have shown correlational evidence, it remains challenging to establish whether the SFS is directly involved in evidence accumulation or whether its observed activity reflects upstream or downstream support processes (e.g., attention or working-memory maintenance) rather than the accumulation computation per se. In this context, it is conceivable to imagine that the SFS would only play a role in a broader network that is not specific to the evidence accumulation process. Moreover, causal evidence from human studies is difficult to obtain because most prior causal studies were only limited to one type of choice or performed in animals, and it is unclear if findings can be translated from animal models to human decision-making. However, disruption of human left SFS with non-invasive stimulation has been shown to impact behavioural performance and response speed in a dynamic face-house classification task, in a manner consistent with a reduction of evidence accumulation during decision-making. However, the domain-specificity of SFS contribution is unclear. Some studies have shown that dlPFC activity may reflect value-based evidence integration (Basten et al., 2010; Sokol-Hessner et al., 2012), suggesting the domain-generality of prefrontal function (Owen, 1997; Petrides, 2005). However, it is hard to directly compare the implicated neural processes to those that underlie perceptual decision-making processes, due to major differences in the stimuli and experimental approaches classically used in each domain (Gold and Shadlen, 2007; Heekeren et al., 2004), and that direct and principled comparisons with other decision-making domains, in general, are largely missing.

Here, we test the domain-specificity of the left SFS, and address the crucial double-dissociation gap in the literature by applying continuous theta-burst transcranial magnetic stimulation (cTBS) followed by functional magnetic resonance imaging (fMRI) while human participants alternated between matched perceptual and value-based choices (Polanía et al., 2014, 2015). We modelled the observed behavioural changes with the DDM, allowing us to causally associate the stimulated SFS region to specific underlying latent subprocesses of the unfolding decision (Mulder et al., 2014; Polania et al., 2015) as well as BOLD activation. Thus, this common evidence accumulation framework provides us with clear testable hypotheses regarding possible effect patterns across behavioural, computational, and neural levels.

Results

The experiment

We recorded functional magnetic resonance imaging (fMRI) data from hungry, healthy participants (n = 20) performing perceptual- and value-based choice-tasks in alternation (Methods and Fig. 1b). For perceptual decisions, participants chose the larger food item, while for value-based decisions, participants chose the food item that they would preferably receive and consume by the end of the experiment. The stimuli and motor responses were identical for both tasks, as in previous experiments (Polanía et al., 2014, 2015). Choice pairings were predetermined based on participant’s individual subjective perceptual- and value-based ratings of the food items, obtained just prior to the scanning session. Perceptual evidence was defined as the size difference (SD) between the food items, whereas value evidence was defined as the difference in value ratings (VD) between the choice alternatives (see Methods and Fig. 1b). A choice was classified as correct when it was consistent with the previously acquired ratings regarding size and preference respectively, i.e., when the larger-rated item was chosen for perceptual decisions or the higher-valued item was chosen for value-based decisions (Polanía et al., 2014, 2015). Our sample size is well within acceptable range, similar to that of previous TMS studies (Philiastides et al., 2011; Rahnev et al., 2016; Jackson et al., 2021; van der Plas et al., 2021; Murd et al., 2021).

Behavioural food choice paradigm, theta-burst stimulation protocol, and behavioural regressions.

(a) Example of decision stage. Participants were cued in advance about the type of decision required. Perceptual decisions required participants to choose the food item with the largest size while value-based decisions required participants to choose the food item they preferred to consume at the end of the experiment. Participants alternated between blocks of perceptual (blue) or value-based (red) choice trials (7-9 trials per task-block). (b) Logistic regression results show that the larger the evidence strength, the more likely decision makers will respond accurately. Choice accuracy is only related to the evidence that is currently task-relevant (size difference SD for perceptual or value difference VD for value-based choice), not to the task-irrelevant evidence (RT is reaction time of current choice). (c) Similarly, our linear regressions show that RTs are negatively associated only with the task-relevant evidence (and lower for perceptual choices overall, captured by regressor CH (1 = perceptual, 0 = value-based)). Consistent with previous findings, the results in (b) and (c) confirm that our paradigm can distinguish and compare evidence processing for matched perceptual- and value-based decisions. Error bars in (b) and (c) represent the 95% confidence interval range of the estimated effect sizes. * p < 0.05, ** p < 0.01, and *** p < 0.001. (d) Theta-burst stimulation protocol. After the fourth pre-TMS run, participants received continuous theta-burst stimulation (cTBS) over the left SFS region of interest (ROI) (area encircled and colored blue). cTBS consisted of 200 trains of 600 pulses of 5 Hz frequency for 50 s.

Our experiment was divided into pre- and post-stimulation blocks. After participants had performed four pre-stimulation session-blocks inside the scanner, they received continuous theta burst stimulation (cTBS) (Huang et al., 2005; Di Lazzaro et al., 2005, 2008) over the left SFS (MNI coordinates, x = −24, y = 24, z = 36; Heekeren et al., 2004; Philiastides et al., 2011; Grueschow et al., 2018). Following this intervention, participants completed four post-stimulation fMRI blocks. By comparing the effects of stimulation on both types of behaviour and brain activity between post- and pre-stimulation blocks, we identify the role of SFS for either type of decision making. In particular, we examined whether the SFS is indeed selectively involved in perceptual decisions as previously suggested (Heekeren et al., 2004, 2006; Rahnev et al., 2016; Philiastides et al., 2011).

Our initial fMRI analyses were conducted at two levels. The first analysis aimed to broadly identify brain areas recruited for each choice task and those common to both. In this analysis, we assessed the average BOLD activity at the task level (perceptual versus value-based), irrespective of evidence accumulation. The second analysis focused on areas representing evidence accumulation specific to each type of choice. Here, we assessed how BOLD activity is modulated by trial-by-trial evidence strength. For this analysis, we used evidence strength from each choice task (perceptual or value-based) as a parametric modulator and regressed trial-by-trial evidence strength with BOLD (see Methods).

Study hypotheses

Previous studies have identified the causal mechanistic role of the SFS in evidence accumulation during perceptual decision-making (Heekeren et al., 2004, 2006; Rahnev et al., 2016), and its effect on stimulation may arise from one of either two channels. That is, one study reported that SFS disruption during a speeded perceptual categorisation task reduced accuracy and increased response times (Philiastides et al., 2011) and found associated decreases in drift rate, the DDM parameter describing the efficiency of sensory evidence integration. In contrast, another human brain stimulation study suggested that behavioural changes due to SFS disruption during a perceptual two-alternative-forced-choice (2AFC) task reflect decreased in the decision threshold, characterised by faster response speed but decreased choice precision. Simulations with the same DDM modelling framework (Rahnev et al., 2016) suggested that the decision threshold parameter could account for individual behavioural changes. Simultaneously acquired fMRI data suggested that SFS does not code the rate of integration but rather the necessary amount of evidence to be accumulated for the perceptual choice at hand (Rahnev et al., 2016).

To this end, we hypothesise that if SFS neurons indeed selectively accumulate perceptual evidence, we should find that their inhibition by cTBS leads to decreases in choice precision and increases in reaction times, a behavioural pattern that corresponds to a decrease in the DDM drift-rate parameter, and to concurrent increases in BOLD signals (caused by prolonged neural evidence accumulation; Fig. 2a-c). Critically, a different pattern can be expected when SFS neurons are involved in setting the criterion, i.e., determining the amount of evidence that needs to be accumulated for a perceptual choice to be taken. In this case, SFS inhibition should result in decreases in both choice precision and reaction times, a decrease in the DDM boundary parameter (Rahnev et al., 2016), and a reduction in associated neural activity due to the lower amount of evidence accumulated during the shorter response time (Fig. 2d-f). Here we directly test these two contrasting scenarios, by characterising the behavioural, neural, and neuro-computational consequences of cTBS to the left superior frontal sulcus (SFS). Crucially, we also investigate for both possible outcomes whether the functional contribution of the SFS during decision making is indeed specific for perceptual choices, by comparing the results between the two matched types of choices.

Study hypotheses.

Scenario 1: left SFS is causally involved in evidence accumulation. Theta-burst induced inhibition of left SFS should lead to reduced evidence accumulation (a), expressed as lower accuracy (a, 2nd row, left), slowing of RTs (a, 2nd row, right), and a reduction of DDM drift rate (b, right) without any effect on the boundary parameter (b, left). Since the neural activity devoted to evidence accumulation (area under the curve) should increase (c, left), we would expect higher BOLD signal in this case (c, right). Scenario 2: left SFS is causally involved in setting the choice criterion. Theta-burst induced inhibition of left SFS should lead to a lower choice criterion (d), expressed as lower choice accuracy (d, 2nd row, left), faster RTs (d, 2nd row, right), and a reduced DDM decision boundary parameter (e, left) without any effect on the DDM drift-rate (e, right). At the neural level, we should observe reduced BOLD activity due to the lower amount of evidence processed by the neurons (f, right), and reflected by the smaller area under the evidence-accumulation curve when it reaches the lower boundary (f, left).

Behaviour: validity of task-relevant pre-requisites

Before scrutinising the role of the left SFS for either type of choice, we first behaviourally and neurally confirmed the validity of our task paradigm. To establish a fair comparison between perceptual (PDM) and value-based decision-making (VDM), we must necessarily show that we can distinctly identify the brain regions associated for each type of choice, and that behaviour is systematically a function of their respective evidence measures. Initial visual inspection shows that choice accuracy/consistency systematically increases (Fig. 2a) and RTs become faster (Fig. 2b) the larger the evidence difference, and this holds across tasks and stimulation conditions. All behavioural and fMRI analyses were performed on valid trials only (see Methods for inclusion criteria). Behavioural regressions confirmed that our task design allowed for a clear computational separation of both choice types: during perceptual decisions, participants relied exclusively on perceptual evidence, as reflected in both increased choice accuracy (main effect SD, β = 0.560, p < 0.001 and VD, β = 0.023, p = 0.178; Fig. 1b) and faster reaction times (RTs) with larger perceptual evidence, but not value-based evidence (main effect SD, β = −0.057, p < 0.001 and VD, β = 0.002, p = 0.281; Fig. 1c). Conversely, participants relied only on value evidence during VDM, as evident from both choice consistency (main effect VD, β = 0.249, p < 0.001 and SD, β = 0.005, p = 0.826; Fig. 1b) and RTs (main effect VD, β = −0.016, p = 0.011 and SD, β = −0.003, p = 0.419; Fig. 1c) irrespective of the items’ size difference. Thus, our results replicate previous findings obtained with a similar paradigm (Polanía et al., 2014, 2015; Grueschow et al., 2015) showing that participants can use exclusively task-relevant evidence to make choices, and they confirm the suitability of our paradigm for directly comparing perceptual and value-based decisions with matched stimuli and motor responses. Across sessions, RTs tended to shorten in both tasks. In line with the HDDM results— selective boundary reductions for PDM and selective nDT shortening for VDM—we interpret the VDM RT speed-ups as reflecting more efficient non-decision (sensorimotor) components rather than changes in evidence accumulation. To illustrate baseline trends, we provide session-wise RT trajectories in the Supplement (see Supplementary Fig. 11 for RT-by-session). For completeness, group-mean accuracies by task are provided descriptively in Fig. 3a; inferential tests focus on evidence-specific effects and TMS-induced changes within task.

Theta-burst stimulation over the left SFS affects choice behavior and selectively lowers the decision boundary for perceptual but not value-based choices.

(a) Choice accuracies/ consistencies and (b) response times (RTs) for perceptual (blue) and value-based (orange) decisions for different evidence levels during pre-cTBS (dark) and post-cTBS (light) stimulation periods. Error bars in (a) and (b) represent s.e.m. Consistent with previous findings, stronger evidence leads to more accurate choices and faster RTs in both types of decisions. Importantly, theta-burst stimulation significantly lowered choice accuracy selectively for perceptual, not value-based decisions (negative main stimulation effect for perceptual decisions and negative stimulation × task interaction; Supplementary Fig. 2c and see also Supplementary Fig. 2a for changes in choice accuracy across runs). Additionally, theta-burst stimulation also significantly lowered RTs in both choice types (negative main stimulation effect; Supplementary Fig. 2c and see also Supplementary Fig. 2b for changes in RTs across runs). (c) Theta-burst stimulation selectively decreased the decision boundary in perceptual decisions only (difference between estimated posterior population distributions; see Methods and Supplementary Fig. 5a for a detailed post-hoc analysis). All the other parameters, particularly (d) the drift rate (see also Supplementary Fig. 5b for post-hoc analysis) remain unaffected by stimulation. Error bars in (c) and (d) represent the 95% confidence interval range of the posterior estimates of the DDM parameters. * p < 0.05, ** p < 0.01, and *** p < 0.001.

fMRI: VDM and PDM distinctly recruit brain processes, while recruiting similar visual and motor processes

In line with the behavioural results that participants depended on different evidence for the two types of choices, initial fMRI analysis revealed that neural activations strongly differed between choice types, despite the fact that participants saw the same images and gave the same motor responses.

We performed two levels of fMRI analyses. The first analysis examined average BOLD activity at the task level, contrasting perceptual versus value-based decisions, irrespective of whether the differences were driven by evidence accumulation or other cognitive processes. This initial analysis aimed to broadly identify brain areas recruited for each choice task and those common to both. First, we found visual and motor areas were jointly activated for both types of choices (p < 0.05, FWE-corrected with cluster forming thresholds at T(19) > 2.9; Supplementary Fig. 1a and Supplementary Table 1). Second, PDM led to stronger recruitment of the posterior parietal cortex whereas VDM led to stronger activations of the medial prefrontal cortex and posterior cingulate cortex (Supplementary Fig. 1b and Supplementary Table 2), all in line with previous findings (Grueschow et al., 2015).

Notably, the left SFS does not yet appear in this contrast analysis. This is because this initial analysis did not include the parametric modulator for evidence accumulation. At this stage, we were evaluating whether the brain can flexibly recruit distinct brain regions based on task design, irrespective of whether these differences arise from variations in evidence accumulation or other cognitive processes. The second analysis assessed how BOLD activity is modulated by trial-by-trial evidence strength. Here, we identified a parametric modulator for evidence accumulation and regressed it on BOLD activity using a general linear model. This analysis revealed that the left SFS shows significant activation specifically when perceptual evidence is parametrically modulated. BOLD activity in the left SFS is detectable only when perceptual evidence is considered6. In fact, previous studies have similarly shown that the SFS only appears once a variable that measures the degree of evidence accumulation is included in the analysis6. Thus, these choice-type-specific brain activations, in response to identical visual input and motor output, ascertain that participants recruit task-specific brain regions depending on the choice domain.

Behaviour: theta-burst stimulation reduces choice accuracy for perceptual decisions only

Our results support the hypothesis that the SFS has a specific role for perceptual decision-making, on several experimental levels. Using a differences-in-differences (DID) logistic regression (Methods), we found that SFS-cTBS led to a significant decrease from pre- to post-cTBS blocks in accuracy for PDM (main stimulation effect, β = −0.465 ± 0.342, p = 0.008; Fig. 3a and Supplementary Fig. 2a), while VDM choice consistency remained unaffected by SFS stimulation (β = −0.042 ± 0.205, p = 0.691; Fig. 3a and Supplementary Fig. 2a). These differences were significant in direct comparison (stimulation × task interaction, β = −0.094 ± 0.087, p = 0.034; Fig. 3a; Supplementary Fig. 2c and Supplementary Table 2) fatigue or habituation effects after checking that the average accuracies in PDM were actually recovering in the second post-stimulation session while there was no change in choice consistency at all during VDM (Supplementary Fig. 2a). Interestingly, our DID linear regression (Methods) revealed that SFS-cTBS had comparable effects on reaction times in both tasks: faster RTs were observed after SFS-cTBS for both PDM (main stimulation effect, β = −0.116 ± 0.067, p = 0.003; Fig. 3b and Supplementary Fig. 2b) and VDM (main stimulation effect, β = −0.125 ± 0.063, p = 0.001; Fig. 3b and Supplementary Fig. 2b), with no significant difference between these two effects (stimulation × task interaction, β = 0.009 ± 0.069, p = 0.795; Fig. 3b; Supplementary Fig. 2c and Supplementary Table 2). Overall, the specific changes in choice accuracy indeed reflect cTBS disruption in left SFS in perceptual decisions. At the same time, the common changes in RTs from the first to the second half of the experiment may not necessarily reflect TMS-related changes in SFS function but rather general training effects common to both tasks (Mawase et al., 2018), but this possibility can only be examined in more detail with computational modelling.

Modelling: SFS-TMS reduces decision boundary only for perceptual decisions

To examine in detail which specific latent decision process was affected by SFS-cTBS, we fit the hierarchical drift diffusion model (HDDM) simultaneously to the accuracy and RT data of our participants. This canonical model of choices allowed us to identify and disentangle the effect of stimulation on various latent variables representing distinct components of the choice mechanism (Ratcliff and Smith, 2004; Ratcliff and McKoon, 2008; Polania et al., 2015; Supplementary Fig. 3 and see Methods).

To investigate the underlying processes through which the cTBS stimulation induced the observed behavioural changes, we fitted a hierarchical Bayesian drift-diffusion model (HDDM; see Methods) simultaneously to the accuracy and RT data of our participants (Fig. 3b-d). Critically, we used the DDM parameters to identify and disentangle the effect of stimulation on choice accuracy from that on RTs (Ratcliff and Smith, 2004; Ratcliff and McKoon, 2008; Polania et al., 2015; Supplementary Fig. 3 and see Methods). A specific focus of this analysis was on whether SFS-cTBS would change the way participants set the choice criterion (decision threshold; Rahnev et al., 2016; Bogacz et al., 2010; Domenech and Dreher, 2010; Herz et al., 2016) or the efficiency with which choice-relevant evidence is accumulated (drift-rate, Philiastides et al., 2011; Basten et al., 2010) (see Methods for more details and Fig. 1b,e). We found that theta-burst stimulation selectively reduced the decision boundary in perceptual decision making (PDM) (see Methods; pmcmc = 0.003; Fig. 3c and Supplementary Fig. 5a), while leaving the decision-relevant parameters, including the drift rate, unchanged (pmcmc = 0.822 for drift rate; Fig. 3d and Supplementary Fig. 5b,c). For VDM, by contrast, no effect of cTBS was observed for either of the two decision-relevant parameters (pmcmc = 0.115 for boundary and pmcmc = 0.758 for drift rate; Fig. 3c,d and Supplementary Fig. 5a,b,c), supporting the specificity of the SFS involvement in perceptual decisions. Full posterior summaries are provided in Supplementary Tables 8–11, and model adequacy is confirmed by posterior-predictive checks of accuracies and RT distributions (Supplementary Figs. 11–17).However, we found that non-decision time (nDT) was selectively reduced in VDM. Overall, our findings indicate that the left SFS is causally involved in modulating the decision threshold. This conclusion was further corroborated by direct comparison of these effects, which showed that SFS-cTBS had a significantly stronger impact on the boundary parameter for PDM compared to VDM (stimulation × task interaction for the decision threshold, pmcmc = 0.045; Supplementary Fig. 5a; there were no such differences for drift-rate; pmcmc = 0.685; Supplementary Fig. 5b).

Modelling: faster RTs during value-based decisions is related to non-decision-related sensorimotor processes

To address the underlying latent process driving RT effects in both choices, we examined other DDM parameters and measurements. The DDM assumes that RTs can be disentangled into a non-decision-related (nDT) component as well as decision times (DT). The nDT is a DDM parameter that indexes constant latencies associated with sensory and motor preparation processes that are invariant across trials with different choice evidence (Verdonck and Tuerlinckx, 2016; Starns and Ma, 2018); in other words, this parameter forms no part of the evidence accumulation process (Feltgen and Daunizeau, 2020; White et al., 2018) and may therefore reflect task learning processes from movement repetition (Mawase et al., 2018). In contrast, decision times is the component of RT where evidence accumulation actually takes place, and we can measure and derive DT using the evidence-dependent DDM parameters (see Methods for more details).

Our results showed that the faster RTs observed for value-based decisions after the stimulation indeed did not reflect evidence-dependent choice processes, but rather a change in non-decision-related sensorimotor processes (nDT) (see Methods; Supplementary Fig. 3): this parameter was decreased after stimulation for VDM (pMCMC = 0.062) but not PDM (pmcmc = 0.707) (Supplementary Fig. 4b and 5c), with a significant difference between these effects (pmcmc = 0.041; Supplementary Fig. 5c). In contrast, estimated decision times was smaller after stimulation during PDM (pmcmc = 0.003; Supplementary Fig. 6a, left), but not VDM (pmcmc = 0.100; Supplementary Fig. 6a, right). Taken together, these results suggest that the simultaneous change in RT reveal completely different computational processes, whereby faster RTs during value-based choice is simply a by-product of task-related learning that may perhaps be unrelated to stimulation, while faster RTs during perceptual choice is actually related to decision-relevant, evidence-dependent latent choice processes. However, completely ascertaining whether such effect from stimulation is due to SFS inhibition, we need clear causal evidence of changes from neural processing. The pattern of non-decision time and decision-time changes is consistent with the posterior-predictive fits shown in Supplementary Figs. 11–17, with numerical posterior summaries in Supplementary Tables 8–11.

fMRI: SFS activation changes for perceptual choices in line with model predictions

To investigate whether our behavioural and computational results directly relate to task-specific disruption of neural activity in left SFS, we investigated BOLD response changes in this brain area after stimulation. We exploited the fact that our fitted DDM and its latent parameters make clear predictions about how BOLD responses in this area should change if the stimulation affects the neural computations involved in setting the boundary for the necessary amount of evidence accumulation. Importantly, these predictions translate to clear parametric regressors that we can use for trialwise analysis of fMRI data (Basten et al., 2010; Domenech et al., 2017; Liu and Pleskac, 2011). More specifically, we expected that the BOLD signal level is proportional to the DDM’s accumulated evidence (aE), defined as the area below the modelled evidence accumulation curve up until the accumulator reaches the decision boundary (Liu and Pleskac, 2011; Domenech et al., 2017; Basten et al., 2010). Using subject-wise DDM latent parameters, the average area below the decision boundary for each evidence level can be computed as a function of each participant’s decision boundary divided by the mean drift rate (see Fig. 1c and 1f and Methods for more details). Using the more detailed trialwise measures, however, the same area can be computed as a function of each trial’s RTs divided by the evidence level, since according to the DDM, the duration of response times is directly proportional to the decision boundary, and the evidence level is directly proportional to the slope of the drift rate (Ratcliff and Rouder, 1998; Ratcliff and McKoon, 2008; see Methods for more details). Exploiting these two known facts from the DDM thus allows us to extend our test of the stimulation effect from individual-specific latent parameters to trialwise regressors and behavioural measures. Higher SFS BOLD signals are associated with higher aE and vice versa (Basten et al., 2010; Liu and Pleskac, 2011; Filimon et al, 2013; Tosoni et al., 2008), implying that a TMS intervention lowering the decision boundary should lower aE and therefore BOLD signals. Crucially, these latent changes predicted by the DDM should also be reflected in the subject-level simulations of accumulated evidence constructed from the DDM parameters.

Thus, we first tested whether our neural hypotheses would already be evident in the simulated trial-wise aE regressors. We used individual parameters identified by fitting our computational framework to simulate expected neural activity on a trial-wise basis across participants. To this end, we derived the predicted aE from the model parameters for each participant. A comparison across cTBS and task conditions confirmed the predicted cTBS-related decrease in accumulated perceptual evidence for PDM (pmcmc = 0.003; Fig. 4a and Supplementary Fig. 6b), the corresponding null effect for VDM (pmcmc = 0.100; Fig. 4b and Supplementary Fig. 6b), and a significant difference for this effect between both choice types (one-sided pmcmc = 0.048; Supplementary Fig. 6b).

Neural representation of accumulated evidence in the left SFS is disrupted after theta-burst stimulation, and is linked with behavior and neural computation.

(a) Left panel: Accumulated evidence (AE) simulation derived from the fitted DDM (left panel). Previous studies have illustrated how the accumulation-to-bound process convolved with the hemodynamic response function (HRF) results in BOLD signals; hence, the simulated AE provides a suitable prediction of BOLD responses in brain regions involved in evidence accumulation. Theta-burst stimulation selectively decreased AE for (a) perceptual (blue), not (b) value-based (orange) decisions (see Supplementary Fig. 6b for post-hoc analysis). We constructed a trialwise measure of accumulated evidence using RTs and evidence strength for our parametric modulator (see Methods). Individual ROIs extracted from the left SFS representing accumulated evidence across runs (right panels; see Methods) show that consistent with the DDM prediction, theta-burst stimulation selectively decreased BOLD response representing AE in left SFS during perceptual, not value-based decisions. Error bars in the left panels of (a) and (b) represent the 95% confidence interval range of the posterior estimates of the DDM parameters, while error bars in their respective right panels represent s.e.m. (c) Post–pre contrasts for the trialwise accumulated-evidence regressor show reduced left-SFS BOLD during perceptual decisions (green overlay), with a significantly stronger reduction for perceptual vs value-based decisions (blue overlay). No reduction is observed for value-based decisions. (d) To test the link between neural and behavioural effects of TMS, regression results show that after stimulation, BOLD changes in left SFS are associated with lower choice accuracy (left panel) for perceptual (PDM, blue) (negative left SFS × stimulation interaction) but not value-based choices (VDM, red), with significant differences between the effects on both choice types (difference-in-difference, DID, green, negative left SFS × stimulation × task interaction). On the other hand, cTBA-induced changes in left SFS activity are unrelated to changes in RT (right panel). Error bars in (d) represent the 95% confidence interval range of the estimated effect sizes. * p < 0.05, ** p < 0.01, and *** p < 0.001. (e) To test the link between neural activity and DDM computations, we included trialwise beta estimates of left-SFS BOLD signals as inputs to the DDM. Alternative models tested whether trialwise left-SFS (LD) activity modulates the decision boundary (α) (Model 1), the drift rate (β), or a combination of both (Models 3 and 4, see Methods and Supplementary Fig. 8 for more details). Model comparisons using the deviance information criterion (DIC, smaller values mean better fits) showed that Model 1 fits the data best, confirming that the left SFS is involved in selectively changing the decision boundary for perceptual decisions.

In the next step, we used the trial-by-trial accumulated evidence as a regressor in the statistical analysis of the BOLD signals, allowing us to test whether the left SFS shows the predicted changes in neural response to varying levels of perceptual evidence. First, we tested whether our predictor of neural accumulated evidence was represented in BOLD signals of similar task-specific areas as reported previously for PDM in SFS (Heekeren et al., 2004, 2006) and for VDM in vmPFC (De Martino et al, 2013; Grueschow et al., 2015). This was confirmed by the data: During PDM, trialwise aE correlated with BOLD activity in the left SFS (peak at = −21, = 26, = 37; SV < 0.05; Supplementary Fig. 7b and Supplementary Table 3) whereas, critically, no significant BOLD activity in the left SFS was observed during VDM. During VDM, aE related to BOLD activity in the ventromedial prefrontal cortex (vmPFC) (peak at = 3, = 38, = −17; SV < 0.05; Supplementary Fig. 7e) and the nucleus accumbens (peak at x = 9, y = 11, z = −11; < 0.05, FWE-corrected with cluster-forming thresholds at T(19) > 2.9; Supplementary Fig. 7e). For both types of choices, domain-general representations of aE were also evident (see Supplementary Fig. 7 and Supplementary Table 3).

We then tested whether cTBS specifically reduced the neural representation of accumulated perceptual evidence in the left SFS for PDM, as predicted by the behavioural and modeling results. In line with these predictions, comparison of the post − pre trial-aE regressor showed a lower BOLD response in left SFS to the trialwise perceptual evidence during PDM (SV < 0.05; Fig. 4c, green patch). This effect was significantly stronger than the corresponding effect on evidence representations in this area during VDM (SV < 0.05; Fig. 4c, blue patch). No effect was found for VDM alone. This indicates that the TMS effect is specific to the SFS during perceptual decisions, not value-based ones, as supported by the BOLD activity analysis. The reduction in BOLD activity in the left SFS during PDM after TMS is consistent with the DDM prediction of a reduction in the accumulated evidence due to a lower decision boundary. Convergent evidence for the specificity of this effect was provided by an alternative hypothesis-guided region-of-interest (ROI) analysis of the regression weights extracted from an a priori ROI-mask of the SFS (see Methods). This showed lower post-stimulation beta values for the trial-aE regressor during PDM (main stimulation effect, β = −0.153 ± 0.054, p = 0.004; Fig. 4a) but not VDM (main stimulation effect, β = 0.078 ± 0.053, p = 0.140; Fig. 4b) and a significant difference in these effects (stimulation × task interaction, β = −0.232 ± 0.075, p = 0.002; Fig. 4a,b). Thus, the fMRI results show that cTBS of the left SFS indeed affects neural processing in this brain structure selectively during perceptual choices, in a way that is consistent with a lowering of the boundary and less accumulated evidence as predicted by the fitted DDM model. This remarkable convergence between the behavioural, modelling, and fMRI results suggests that the left SFS is indeed causally involved in setting decision criteria for choices based on perceptual evidence, but not based on subjective values.

fMRI and modelling: neural-HDDM shows that perceptual-choice accuracy and boundary setting reflect trial-by-trial changes in SFS activity

If perceptual-decision performance depends specifically on activity in the left SFS, then trial-wise choice accuracy should relate to trial-wise BOLD activity in the SFS during perceptual decisions, over and above the mean effects of evidence level. To test this, we regressed choice accuracy/consistency on trial-by-trial BOLD activity extracted from the left SFS ROI, choice type, and TMS, while controlling for the evidence provided by the stimulus pairs on each trial (see Methods for details). In line with our prediction, we observed that the relation between perceptual-choice accuracy and trial-by-trial SFS activity was significantly decreased by TMS (SFS × stimulation interaction, β = −0.196 ± 0.128, p = 0.003; Fig. 4d), independently of the corresponding effects for choice evidence (SD main effect, β = 0.524 ± 0.082, p < 0.001, VD main effect, β = 0.197 ± 0.012, p = 0.001, SFS × SD interaction, β = −0.041 ± 0.046, p = 0.365, SFS × VD interaction, β = 0.055 ± 0.041, p = 0.183). This effect was clearly specific for PDM, since no such effects were observed for VDM (SFS × stimulation interaction β = 0.099 ± 0.242, p = 0.422; SFS × stimulation × task interaction, β = −0.072 ± 0.051, p = 0.005; Fig. 4d) and for RTs during both types of choices (SFS × stimulation interaction, perceptual: β = −0.031 ± 0.053, p = 0.367; Fig. 4d; accuracy: SFS × stimulation interaction, β = −0.012 ± 0.050, p = 0.650; Fig 4d).

We further investigated whether the relation between trialwise SFS activity and choice outcome indeed reflected an SFS role for perceptual boundary setting, as suggested by the DDM results presented above. To confirm this neurally, we set up several DDMs with trialwise SFS activity as an additional modulator for DDM parameters (on top of choice evidence; see Methods and Herz et al., 2016, 2017; Turner et al., 2015). More specifically, we tested several neural-DDMs in which trialwise SFS activity either modulated the decision threshold only (Model 1; Supplementary Fig. 8a), the drift rate only (Model 2; Supplementary Fig. 8b), or both parameters separately (Model 3; Supplementary Fig. 8c) or jointly (Model 4; Supplementary Fig. 8d). We compared these neural HDDMs to our baseline HDDM without neural inputs (see Methods for more details and Supplementary Fig. 3), allowing us to test across all conditions and choice types whether model evidence was enhanced when adding a potential trial-by-trial influences of SFS activity to the experimental inputs. Thus, the reported model evidence criterion (DIC) provides an additional formal test of whether the cTBS-influenced SFS activity relates selectively to the decrease of the decision boundary for perceptual choices only. Consistent with this prediction, Model 1 where SFS activity modulated the decision threshold only, outperformed all other models and model evidence showed improvements versus the baseline model (relative DIC = −28.65; Fig. 4b). These results provide direct evidence that neural computations in the left SFS support criterion setting for perceptual evidence accumulation.

fMRI and connectivity: TMS affects SFS functional connectivity during perceptual choices

Overall, our results clearly indicate that cTBS to the left SFS disrupts selectively a neural process related to setting the criterion for perceptual evidence accumulation. At this point, we consider the possibility that cTBS may conceivably change the functional communication of the SFS with other brain areas involved in initial processing of the perceptual information necessary to make a choice. We explored this possibility by investigating whether cTBS affected functional coupling of the SFS. A psychophysiological interaction (PPI) analysis seeded in left SFS and modulated by aE indeed revealed stronger coupling with occipital cortex (OCC) after cTBS (peak at x = −28, y = −85, z = −2; p < 0.01 FWE-cluster-forming thresholds at T(19) > 2.9; Fig. 5a). Interestingly, the activity peak in visual cortex showing evidence-dependent coupling with SFS, overlaps with the spatiotopic neural representation of the stimulus items in the visual field during decision making. We identified this overlap using a conjunction analysis of the PPI result and a contrast regressing BOLD signal on trial-by-trial stimulus onsets of both choice types (at familywise-error-corrected thresholds). Moreover, we used the latter contrast to define fully independent regions-of-interest (ROIs) in occipital cortex processing the visual stimuli independent of task type and performed an ROI analysis on the individual SFS-OCC-PPI betas extracted for each participant. This confirmed that evidence-related functional coupling is increased by stimulation during PDM (main stimulation effect, β = 0.330 ± 0.284, p = 0.022) but not VDM (main stimulation effect, β = −0.186 ± 0.247, p = 0.139; Fig. 5b; stimulation × task interaction, β = 0.517 ± 0.44, p = 0.021; Fig. 5a). Thus, our exploratory analysis shows that cTBS to the left SFS leads to stronger functional coupling with occipital areas involved in processing the visual stimuli, perhaps consistent with increased downstream demand on visual-related resources when upstream evidence accumulation regions are impaired.

SFS-TMS-related changes in behaviour and neural computations are accompanied by increased functional coupling between the left SFS and occipital cortex.

(a) Psychophysiological interaction (PPI) analysis reveals an area in occipital cortex showing increased functional coupling with the left SFS during perceptual choices. (b) ROI analysis of individual PPI betas shows that aE-related functional coupling between the left SFS and OCC is selectively increased post stimulation during perceptual (left panel) but not value-based decisions (right panel). Error bars in (b) represent s.e.m. (c) Regression results testing the link between cTBS effects on left SFS-OCC functional coupling and behaviour. Increased SFS-OCC coupling is associated with lower choice accuracy (left panel) specifically for perceptual (PDM, blue, negative OCC × stimulation interaction) but not value-based choices (VDM, red). In addition, increased functional coupling is also associated with faster RTs (right panel) for perceptual (blue, negative OCC × stimulation interaction) and slower RTs for value-based choice (red, positive OCC × stimulation interaction). Error bars in (c) represent the 95% confidence interval range of the estimated effect sizes. * p < 0.05, ** p < 0.01, and *** p < 0.001.

We further explored whether this TMS-induced increase in functional coupling between the left SFS and OCC is related to changes in behaviour and specific neural computations during perceptual decisions. To test this, we related these effects to individual measures of choice behaviour and latent DDM parameters for each participant. This revealed that stimulation-induced increases in SFS-OCC coupling were associated with lower accuracy (OCC × stimulation × task interaction, β = −0.225 ± 0.142, p = 0.002; Fig. 5c) and shorter RTs (OCC × stimulation × task interaction, β = −0.325 ± 0.238, p = 0.007; Fig. 5c) for PDM, but not VDM. Taken together, these results show that the causal behavioural and computational changes during perceptual decision-making due to left SFS-TMS may relate not just to local neural changes in SFS, but also to the way this brain structure communicates with visual cortex.

Discussion

Our study shows that the left SFS serves a domain-specific causal role in the accumulation of perceptual evidence, and particularly, the underlying computations affected by non-invasive brain stimulation is the setting of the choice criterion during perceptual, but not value-based choice. Our findings are more in line with that of Rahnev et al. (2016), who also suggested that SFS disruption leads to lower threshold setting. In contrast, previous work has also shown that SFS activity correlates with the evidence accumulation process reflected by the drift rate. These findings are not necessarily incompatible, but rather, they may reflect the interlinked computational mechanisms of the SFS in perceptual evidence accumulation, where the SFS plays a role in setting the decision threshold. Furthermore, our findings, in a way, contribute in closing the double dissociation gap left by previous work (Polanía et al., 2015), where it was shown that frontoparietal connectivity is causally specific to the precision of value-based, not perceptual choice using the same matched perceptual- and value-based choice task. More importantly, our study provided results above and beyond current standards of causal studies in humans (Philiastides et al., 2011; Rahnev et al., 2016; Polanía et al., 2015; Murd et al., 2020), that was only observed in animal studies (Ding and Gold, 2012b; Erlich et al., 2015; Hanks et al., 2015). Here, we simultaneously showed that causal TMS effects affected behaviour, latent computations, and more crucially, neural circuitry, as observed by changes in fMRI-BOLD activation after stimulation.

Earlier version: Many human decision neuroscience studies have employed model-based approaches to identify BOLD signals that correspond to computational processes (Forstmann et al., 2011; Palmeri et al., 2017; Wijeakumar et al., 2017). However, the links between neural and latent computational processing established by these studies is largely correlational (Logothetis, 2008; Poldrack, 2006; Ramsey et al., 2010), and there are many model alternatives that could possibly account for BOLD signals. Our study illustrates that causal manipulations induced by targeted functional inhibition of brain areas can provide decisive information and provide more direct support for neurocomputational mechanisms posited by cognitive models. Specifically, our study underlines that the DDM provides a plausible mechanistic account of the decision process (Herz et al., 2016, 2017; Turner et al., 2015), by showing that left SFS inhibition by cTBS affects the evidence representation posited by the model consistently across behavioural, computational, neural, neural-behavioural levels. Importantly, our results directly link changes in behaviour to changes in both latent computations and neural processing, by demonstrating how raw trialwise neural signals from the left SFS can augment the DDM to explain behaviour. This suggests that once brain stimulation studies have established (causal) correspondence between neural activity and latent variables in decision models, such models can be fruitfully extended by neural measures to provide a more complete characterisation and prediction of choice behaviour and potentially its malfunctions. Our fMRI inferences rest on model-based assumptions linking accumulated evidence to BOLD amplitude. Alternative mechanisms—such as time-dependent (collapsing) boundaries—could attenuate the prediction that weaker-evidence trials yield longer accumulation and larger BOLD signals. While our behavioural and neural results converge under the DDM framework, we acknowledge this as a general limitation of model-based fMRI. The within-participant design enhances statistical sensitivity, yet the absence of an a priori power analysis constrains our ability to rule out small effects, particularly for null results in VDM.

Revised version: Many human decision neuroscience studies have employed model-based approaches to identify BOLD signals that correspond to computational processes (Forstmann et al., 2011; Palmeri et al., 2017; Wijeakumar et al., 2017). However, the links between neural and latent computational processing established by these studies is largely correlational (Logothetis, 2008; Poldrack, 2006; Ramsey et al., 2010), and there are many model alternatives that could possibly account for BOLD signals. Our study illustrates that causal manipulations induced by targeted functional inhibition of brain areas can provide decisive information and provide more direct support for neurocomputational mechanisms posited by cognitive models. Specifically, our study underlines that the DDM provides a plausible mechanistic account of the decision process (Herz et al., 2016, 2017; Turner et al., 2015), by showing that left SFS inhibition by cTBS affects the evidence representation posited by the model consistently across behavioural, computational, neural, neural-behavioural levels. Importantly, our results directly link changes in behaviour to changes in both latent computations and neural processing, by demonstrating how raw trialwise neural signals from the left SFS can augment the DDM to explain behaviour. This suggests that once brain stimulation studies have established (causal) correspondence between neural activity and latent variables in decision models, such models can be fruitfully extended by neural measures to provide a more complete characterisation and prediction of choice behaviour and potentially its malfunctions.

Added in the revised version: Our fMRI inferences rest on model-based assumptions linking accumulated evidence to BOLD amplitude. Alternative mechanisms—such as time-dependent (collapsing) boundaries—could attenuate the prediction that weaker-evidence trials yield longer accumulation and larger BOLD signals. While our behavioural and neural results converge under the DDM framework, we acknowledge this as a general limitation of model-based fMRI. The within-participant design enhances statistical sensitivity, yet the absence of an a priori power analysis constrains our ability to rule out small effects, particularly for null results in VDM. Consistent with this, model-fit indices and posterior-predictive checks indicated that the HDDM captured perceptual decisions somewhat more tightly than value-based decisions. This pattern likely reflects additional sources of variability in subjective value-based choices, together with our decision to keep the model architecture identical across tasks for principled comparison. Nevertheless, posterior-predictive simulations confirmed that, also for the value-based task, the model reproduced choice proportions and full RT distributions at the group and subject level (Supplementary Figs. 11-16), indicating that the fit quality is sufficient for the present inferences. Thus, although very small stimulation effects on value-based decision parameters cannot be completely excluded, the convergent behavioural, computational, and neural evidence points to a selective impact of cTBS on the perceptual decision process implemented in left SFS.

Specificity of the SFS during perceptual decisions – only in humans?

To this end, our study reinforces already established findings by showing that the causal role of SFS during perceptual decisions (Heekeren et al., 2004, 2006; Philiastides et al., 2011; Rahnev et al., 2016) is a specialised function of evidence integration. Furthermore, our finding of a selective role of left SFS in perceptual evidence accumulation is particularly intriguing. The area appears to be uniquely developed in the human brain, with no close anatomical homologue in other species. In the animal literature, most prefrontal disruption studies in non-human primates have focused on the frontal eye fields (FEF) (Ding and Gold, 2012a; Hanks and Summerfield, 2017; Shadlen and Newsome, 1996) and in rodents on the frontal orienting fields (FOF) (Erlich et al., 2015; Hanks et al., 2015). While we and others observed disruption of the evidence accumulation process after interfering with SFS function in humans (Philiastides et al., 2011; Rahnev et al., 2016), disruption of the FOF in rodents has not affected behaviour at all or in a qualitatively different manner (Brody and Hanks, 2016; Erlich et al., 2015; Hanks et al., 2015). However, the results of electrical stimulation of the FEF in monkeys (Ding and Gold, 2012a; Hanks and Summerfield, 2017) cannot necessarily be directly compared with TMS studies of human SFS, since FEF and SFS in humans are both structurally and functionally distinct (Murd et al., 2020; Rahnev et al., 2016). Thus, while it is tempting to speculate that the SFS perceptual evidence accumulation process identified here may be specific to humans, it is possible that researchers may have to further consider other putative homologues across species that may truly correspond to the SFS area stimulated here (Brunton et al., 2013; Hanks and Summerfield, 2017).

Do value-based decisions also rely on distinct PFC areas?

Added in the revised version: At the same time, our findings do not rule out the possibility of SFS involvement during value-based choice, where it may have a secondary function (but one that does not involve value evidence accumulation) or other specialised roles in decision-making. An alternative interpretation is that cTBS to left SFS induces a general speeding of perceptual processes that are engaged in both tasks. Under this view, the same perturbation of a perceptual decision stage could manifest as a reduction in decision boundary in the pure perceptual task, yet be absorbed as a shortening of non-decision time in the value task when modelled with a single accumulator. Our data cannot completely rule out this multi-stage account for the value task. Importantly, however, such a scenario remains compatible with our central conclusion that SFS plays a primarily perceptual role: any contribution of SFS to the value task would arise via an upstream perceptual component, while the value-comparison process proper is likely implemented in distinct valuation circuitry. The task-specific modulation of BOLD activity in SFS by cTBS during perceptual, but not value-based, choice—together with preserved value sensitivity and the absence of value-related SFS signals—strongly supports this interpretation.

Earlier version: At the same time, our findings do not rule out the possibility of SFS involvement during value-based choice, where it may perhaps have a secondary function (but one that does not involve evidence accumulation) or even other specialised functions in decision-making. For example, previous work has suggested that during value-based decisions, the dlPFC overall interacts with the vmPFC in modulating the value signal to facilitate self-control (Hare et al., 2009). Moreover, value-based decision-making entails a large array of choice types with varying degrees of complexity. For instance, more complex types of value-based decisions entail decisions under risk (Glickman et al., 2019), intertemporal choice (Peters and D’Esposito, 2020), and strategic and social decisions (Hutcherson et al., 2015), which may plausibly recruit the PFC due to working memory demands (Barbey et al., 2013), adjustment of decision time (Sokol-Hessner et al., 2012), or cost-benefit computations (Basten et al., 2010). In light of these many additional decision types based on preferences in the value-based choice domain, the functional specificity of SFS to perceptual decisions we claim here may have to be viewed with caution. Future studies should consider exploring the comparison between more complex types of value-based decisions with perceptual decisions, while taking great care in matching the degree of complexity between the two choice domains to avoid confounds induced by context or task difficulty.

Revised version: Beyond our paradigm, previous work has suggested that during value-based decisions the dlPFC more broadly interacts with the vmPFC to modulate value signals and facilitate self-control (Hare et al., 2009). Moreover, value-based decision-making encompasses a wide range of choice types with varying degrees of complexity. For instance, more complex value-based decisions entail choices under risk (Glickman et al., 2019), intertemporal choice (Peters and D’Esposito, 2020), and strategic or social decisions (Hutcherson et al., 2015), which may plausibly recruit PFC regions due to working memory demands (Barbey et al., 2013), adjustment of decision time (Sokol-Hessner et al., 2012), or cost–benefit computations (Basten et al., 2010). In light of these additional value-based decision types, the functional specificity of SFS to perceptual decisions that we propose here should be viewed with some caution. Future studies should therefore compare more complex value-based decisions with perceptual decisions, while carefully matching task demands and complexity across domains to avoid confounds induced by context or difficulty.

The role of SFS in choice criterion setting during perceptual decisions

That being said, the same holds true with the varying degrees of complexity across perceptual decision-making processes. Our findings show that the mechanism of which the left SFS is causally involved is modulating the decision threshold, with clearly consistent results across behavioural, computational, neural, and neural-behavioural levels. Our findings are more in line with that of Rahnev et al. (2016), who also suggested that SFS disruption leads to lower threshold setting. In contrast, previous work has also shown that SFS activity correlates with the evidence strength in the accumulation process, as reflected by the drift rate (Basten et al., 2010; Heekeren et al., 2004, 2006). In support of this notion, cortical activity disruption with repetitive transcranial magnitude stimulation (rTMS) resulted in lower choice accuracy and slower RTs (Philiastides et al., 2011). However, it is important to note that these results are not at all incompatible, but reflect the interlinked computational mechanisms of the SFS in perceptual evidence accumulation. All these studies, including ours, point out that the DDM’s drift rate and decision boundary are both decision-relevant latent mechanisms with distinct as well as overlapping implications on choice behaviour: impairments in both boundary and drift lead to lower choice consistency (Cavanagh et al., 2011; Green et al., 2012; Herz et al., 2016; Philiastides et al., 2011; Rahnev et al., 2016). The main difference in fact concerns reaction times: lower drift rate implies slower RTs while a lower boundary implies faster RTs. The differences in RT-TMS effects of our study and that of Philiastides et al. (2011) may reveal the differences in task design and the nature of the stimuli. For instance, the study by Philiastides et al. (2016) used dynamic series of briefly presented sequential face-house stimuli, varying the strength of each stimulus within noise to vary the evidence levels. By contrast, our study and that by Rahnev et al. (2016) presented a static stimulus pair simultaneously during a 2AFC task and varied evidence not with noise but by stimulus difference.

We speculate that the goals of these tasks (i.e., discriminability from noisy dynamic stimuli versus stimulus size difference between static stimuli) may sensitively affect different latent processes. Previous studies have shown that different tasks can produce proximally similar behaviour but may involve different goal functions and therefore, computationally distinct processes (Heng et al., 2020). For instance, noisy dynamic stimuli entail the accumulation of sensory evidence until the decision-maker can form representations suitable for choice discrimination (Ratcliff and McKoon, 2008); thus, such stimuli are sensitively modulated by evidence strength, and reflected by a lower drift-rate when SFS is impaired (Philiastides et al., 2011) and where consequently, the circuitry takes more time to form perceptual representations. In contrast, evidence in the form of size difference may depend on the sensitivity to which SFS can detect the difference, and this is modulated by the decision threshold (Herz et al., 2016; Cavanagh et al., 2011; Green et al., 2012, 2013). Here, SFS impairment to accumulate evidence in the context of static, simultaneously presented stimuli may result in the inability of the circuitry to reliably discriminate size differences, resulting in the early termination of the accumulation process, which is behaviourally reflected as lower choice accuracy and faster RTs and computationally as lower thresholds. While a common evidence framework can plausibly reconcile our findings and that of previous studies, future studies should consider addressing this issue directly, by comparing the different goal functions and stimulus displays within perceptual decision-making tasks.

Functional coupling between left SFS and visual cortex

Furthermore, our study further expands our understanding of SFS function vis-à-vis other brain regions. It is well-established that the prefrontal cortex is structurally connected with many other brain regions (Wycoco et al., 2013) and may flexibly interact functionally with different areas depending on choice demands. Our exploratory connectivity results suggest that the SFS role for domain-specific accumulation of perceptual evidence is not just a local phenomenon but extends to functional communication with visual areas. Inhibition of this area’s functional contribution to evidence accumulation led to an increase in its functional coupling with areas in occipital cortex representing the stimuli visually upon which choices were based. The changes in functional coupling strength between the two cortical regions also corresponded to observed behavioural and latent computational changes. This suggests that perceptual choices rely not only on local processing in SFS but on an integrated functional circuit, comprising both SFS and occipital cortex, at least for decisions based on visual stimuli as studied here. Though exploratory, our results are generally consistent with an occipito-frontal information exchange but extend it specifically to the SFS during perceptual evidence accumulation (Bullier et al., 1996).

We can speculate why the occipital cortex may have been recruited after inhibition of the left SFS via cTBS stimulation. For example, it is possible that cTBS-related impairments in the accumulation mechanism implemented by the SFS biases the system to rely on second-best suboptimal mechanisms for solving the tasks, such as template matching from working memory. Previous work has provided converging evidence that maintenance of visual information in working memory enhances coupling between sensory processing in visual cortex and information storage in lateral prefrontal cortex (Gazzaley et al., 2007; Serences et al., 2009). In fact, it has been suggested that SFS is canonically organised in “memory receptive fields” (Postle, 2016) that may be more heavily taxed when direct accumulation mechanisms for sensory input are impaired, as in the case of cTBS manipulations. Of course, there are many other candidate mechanisms, such as attention or working memory, that may be more heavily taxed to compensate for the excitability manipulation of the SFS area specialized for processing the sensory evidence, as suggested by previous work on prefrontal-occipital interactions during various attention and working memory tasks (Zanto et al., 2011). Overall, the exploratory nature of our analysis warrants future investigation on the directionality of information flow between occipital cortex and SFS. Additionally, future studies should also test whether perceptual choices based on other sensory modalities (e.g., touch, audition) lead to a flexible coupling of SFS with the specific sensory areas processing these stimuli. In any case, our study shows clearly that in the healthy, undisrupted human brain, left SFS plays a key role in transforming perceptual evidence into choices.

Implications for theories of PFC organisation

Our study also contributes to our understanding of prefrontal cortex (PFC) functional organisation, given the considerable debates surrounding its organising principles (Owen, 1997). Previous studies have posited that different PFC regions contribute to specific aspects of information processing, in a manner that that can be flexibly applied to all types of information, be it from different sensory modalities or in different cognitive formats (Petrides, 2005). Prevailing perspectives have also proposed an anterior-to-posterior hierarchy in PFC for the purpose of general cognitive control and executive function (Nee and D’Esposito, 2016), suggesting that the main role of the PFC is largely in the domain of higher-order cognitive and abstract operations that transcend specific functional domains (Domenech and Koechlin, 2015; Koechlin and Summerfield, 2007). In contrast, our finding of a domain-specific causal role of SFS in evidence accumulation for perceptual decision-making suggests that the PFC is organised as a collection of fractionated sub-regions, such that each region processes different types of information (Goldman-Rakic and Leung, 2002), as opposed to a systematic hierarchy. Moreover, the fact that the SFS is even involved in the integration of low-level perceptual evidence (Heekeren et al., 2004, 2006, 2008; Philiastides et al., 2011; Rahnev et al., 2016) implies that the PFC’s role is not limited to higher-order cognitive function. Overall, our findings are in no position of proposing an overarching framework of overall PFC organisation, given the limited area in posterior dlPFC targeted by our study, but rather a call to consider alternative views where the underlying organisational principles are more fractionated and less hierarchical.

Implications for computational psychiatry

Finally, our finding of SFS causal involvement in decision-threshold-setting during perceptual decision-making may offer clinical implications. Particularly, manifestations of impulsive behaviour (Heyes et al., 2012) are largely apparent in clinical populations with aberrations in decision threshold setting (Herz et al., 2016). However, most studies of these disorders have focused on impulsive behaviour induced by reward or preferences (Glimcher et al., 2007). It is important to note here that reward impulsivity is only one of the many domains of aberrant behaviour in clinical populations. Perceptual impulsivity is also important, since many of the behavioural and cognitive deficits are closely linked to impairments in perceptual function (Fuermaier et al., 2018). For instance, impulsive behaviour can also arise in non-reward-related settings, such as when perceptually discriminating size differences where less accurate and faster responses have been observed in people with addiction disorders (Banca et al., 2016) and borderline personality disorder (Berlin and Rolls, 2004). Perceptual deficiencies are also prevalent in clinical populations with attention-deficit hyperactivity disorder (ADHD) or Parkinson’s disease, and thought to be linked to impairments in the dopaminergic system (Fuermaier et al., 2018). Prior causal evidence from deep-brain stimulation (DBS), in particular, has shown that disrupting the STN lowered decision thresholds, thus increasing this perceptual impulsivity among Parkinson’s patients (Herz et al., 2016, 2017). Our findings that TMS of the left SFS causally and selectively lowered the decision boundary during perceptual decisions suggest that the lateral prefrontal cortex may be functionally integrated with these cortico-striatal and cortico-subthalamic nuclei (STN) pathways (Bogacz et al., 2010; Green et al., 2012, 2013).

Overall, impulsive behaviour is not exclusive to the reward domain, and our results suggest that there is something to gain from understanding impulsive behaviour in non-reward settings requiring decisions on perceptual information. Maladaptive behaviour may not only reflect individual wants or likings, often assumed by addiction studies, but could also be a function of low-level sensory or higher-order cognitive processes that have so far not fully been accounted for (Fuermaier et al., 2018). This may have serious implications for how cognitive therapies or interventions are designed, and our findings may provide useful insights in guiding such future work. Particularly, it is worth exploring to what degree the left SFS and its connections are structurally or functionally different in clinical populations, and whether these impulsive tendencies can be captured by sequential sampling models, such as the DDM.

Materials and methods

Participants

Twenty healthy right-handed volunteers (ages 20-30; 8 female) with normal or corrected-to-normal vision participated in the study. Participants were fully informed about the features of the experiment. No participant suffered from any neurological or psychological disorder or took medication that interfered with their participation in the study. Participants received monetary compensation for participation and performance of the perceptual choices, as well as one food item to consume after the experiment depending on a random value-based choice trial. The experiments conformed to the Declaration of Helsinki and the experimental protocol was approved by the Ethics Committee of the Canton of Zurich.

Experimental paradigm

We asked participants to refrain from eating for 3 hours before the start of the experiment. Our experiments took place between 0800 and 1900 hr during the day. The experiment consisted of two steps: (1) a rating task outside the scanner and (2) a decision-making task inside the scanner. During the rating task, we asked participants to provide perceptual- and value-based ratings of the same set of 61 food images using an on-screen slider scale. All of the food items were in stock in our lab and participants were informed about this via visual inspection. For perceptual ratings, participants rated—on a scale from 5 to 100 percent in steps of 5 percent— how much of the black background within the white square perimeter was occupied by the food item. For value-based ratings, participants rated—on a scale from 5 to 100 in steps of 5—how much they wanted to eat the presented food item at the end of the experiment. We instructed participants that the midpoint of the scale in value-based ratings indicated indifference.

Earlier version: After rating the food items, an algorithm selected a balanced set of perceptual and value-based trials divided into four evidence levels, E. The evidence levels are based on the absolute difference between the average ratings of the food items paired in each trial. We define perceptual evidence as the absolute size difference between the two food items. On the other hand, we define value-based evidence as the absolute value difference between the two food items. In particular, the evidence levels for perceptual trials, Ep, are: while the evidence levels for value-based trials, Ep, are: where is the average food item from the two ratings while rbiggest − rsmallest and rbest − rworst represent the ratings’ difference for the pairs presented for perceptual and value-based choices, respectively.

Revised version: After rating the food items, an algorithm selected a balanced set of perceptual and value-based trials divided into four evidence levels, E. The evidence levels are based on the absolute difference between the average ratings of the food items paired in each trial. We define perceptual evidence as the absolute size difference between the two food items. On the other hand, we define value-based evidence as the absolute value difference between the two food items. In particular, the evidence levels for perceptual trials, Ep, are: while the evidence levels for value-based trials, Ev, are: where is the average food item from the two ratings while rbiggest − rsmallest and rbest − rworst represent the ratings’ difference for the pairs presented for perceptual and value-based choices, respectively.

Inside the scanner, participants performed the decision-making task for which they chose between two food items, based on whether they were accumulating perceptual or value-based evidence. We matched the visual sensory stimuli of the food items as well as their motor outputs across the two choice types. The only difference was the type of evidence participants had to accumulate to make a choice. Each trial started with presentation of a central fixation marker (length ∼ 0.8°, height ∼ 0.3°). Next, a centrally presented word indicated whether participants would perform a perceptual (word ‘AREA’) or value-based (word ‘LIKE’) choice. On the subsequent screen, the task cue was replaced by either the letter ‘A’ or ‘L’ (∼ 0.2°) to remind participants that they were in a perceptual or value-based block, respectively. Two food items were simultaneously displayed, one above and one below the screen (y eccentricity 3.6°; a white square of 6° width surrounded each food item). Blocks alternated between perceptual and value-based choices in a given session (7-9 trials per task block). Participants pressed one of two buttons on a keypad with their right middle finger (upper item) or right index finger (lower item) to indicate their choice. On a given trial, participants had 3 seconds for their choice; otherwise, the trial would be regarded as a ‘missed trial’ and would not enter the analysis. Analyses were conducted on valid trials only, defined as trials with a registered response within the task’s response window and passing pre-specified validity checks; trials without a response were excluded and not analysed. Participants made correct or consistent choices when they chose the food item with the higher rating as indicated in the double ratings task prior to entering the scanner. After the experiment, participants stayed in the room with the experimenter while they ate the food that was selected based on the participants choice in one randomly selected VDM trial. During perceptual decisions-making blocks, participants were rewarded with 0.5 CHF every time they correctly discriminated between the size difference of the two food items presented on the monitor screen.

The experiment had a total of 256 trials divided into 8 sessions of 32 trials each. The first 4 sessions were pre-stimulation sessions where participants performed the task without stimulation. The last 4 sessions were post-stimulation sessions during which participants performed the choices with decreased neural excitability in the SFS due to the preceding continuous theta-burst stimulation. The 256 trials were fully balanced across all factors (trial type: perceptual or value-based; evidence levels: 1 to 4; correct response: up or down).

Stimulation protocol

We applied continuous theta-burst stimulation (cTBS) (Huang et al., 2005; di Lazzaro et al., 2005, 2008) to exogenously induce cortical inhibition of our region of interest (ROI), an area in the left superior frontal sulcus (SFS) (MNI coordinates: x = −24, y = 24, z = 36) (Heekeren et al., 2004; Philiastides et al., 2011). Before the main fMRI experiment, we identified the stimulation site over the left SFS (MNI coordinates: x = −24, y = 24, z = 36) (Heekeren et al., 2004; Philiastides et al., 2011) based on previous studies and each individual’s stimulation intensity. In an initial fMRI session, we acquired high-resolution T1-weighted 3D fast-field echo anatomical scans used for subsequent neuro-navigation (181 sagittal slices, matrix size = 256 × 256, voxel size = 1 mm3, TR/TE/TI = 8.3/2.26/181 ms, 3T Philips Achieva). The hand area of the left M1 (motor hotspot) was determined by identifying the first dorsal interosseous (FDI) movement-evoked potentials (MEPs) induced by transcranial magnetic stimulation (TMS) pulses. We delivered single monophasic TMS pulses using a figure-of-eight coil attached to the TMS stimulator. We then marked an equidistant circular grid on each individual’s anatomical MRI scan using a neuro-navigation system over the hand motor region, located at the anterior portion of the central sulcus. We localised the optimal motor hotspot as the point in the grid that elicited the strongest FDI MEPs from TMS pulses. Once we selected the motor hotspot, we asked participants to activate their FDI by pressing their thumb and index finger at about 20% maximum force in order to obtain their active motor threshold (AMT). We defined the AMT as the minimal TMS intensity required to produce MEPs of ≥ 200 mV amplitude (measured with Magventure MRi-B91) in ≥ 5–10 consecutive pulses. We retested the AMT by visually inspecting the FDI twitches triggered by TMS pulses over the marked optimal hotspot. The average AMT outside the scanner was 52.35 ± 6.27 percent while the AMT inside the scanner was 52.91 ± 6.18 percent. We applied cTBS at an intensity of 80% of the individual’s AMT. The cTBS protocol contained bursts of 3 pulses at 50 Hz. This protocol has been shown to reduce cortical excitability for at least 30 minutes (Huang et al., 2005). Every burst was repeated at a rate of 5 Hz, resulting in 200 bursts with a total of 600 pulses delivered within 40 seconds.

Before moving our participant into the scanner, we marked the motor hotspot as well as the stimulation site on a swimming cap fixed in position by straps. Participants wore this cap while they were inside the scanner. Before the start of the fifth session, participants received cTBS over the left SFS. We used a figure-of-eight MR-compatible TMS coil (MRi-B91) attached to a TMS stimulator. Participants returned to the scanner after receiving stimulation and proceeded to complete the last four sessions. On average, the post-TMS fMRI task started 228 ± 41 sec after the end of theta-burst stimulation following established protocols from previous studies (Knecht et al., 2003; Philiastides et al., 2011; Thut and Pascual-Leone, 2010). Given the established timeline of cTBS effects (Huang et al., 2005), we expected the stimulation effects to weaken over time due to neural recovery. In line with established procedures, we treated the first two post-stimulation sessions as the actual post-cTBS period and the last two post-stimulation sessions as a recovery period (Philiastides et al., 2011).

Differences-in-differences

Earlier version: We implemented a differences-in-differences (DID) regression analysis to identify the causal relationships based on stimulation-induced neural inhibition in SFS. We used the identical DID for behavioural, computational, neuroimaging and connectivity analyses. Here, we use the following notation: task conditions Task (perceptual, Task = 1; value-based, Task = 0); stimulation conditions TMS (pre-TMS = 0 and post-TMS = 1); V is our variable of interest, which may be behavioural or neural; and, the causal treatment effect, ϕ(V|Task, TMS), takes the following form: where 𝔼(V|Task, TMS) is the expected value of the variable of interest, V, given task and TMS condition. The first difference on the right-hand side captures the average stimulation effect for PDM while the second difference captures the average stimulation effect for VDM. The overall difference assumes that if behaviour will be the same after stimulation, then there is no effect, ϕ = 0 (Angrist and Pischke, 2009; Bertrand et al., 2004). But if there is a stimulation effect and it impairs behaviour or neural activity, then ϕ < 0.

Revised version: We implemented a differences-in-differences (DID) regression analysis to identify the causal relationships based on stimulation-induced neural inhibition in SFS. We used the identical DID for behavioural, computational, neuroimaging and connectivity analyses. Here, we use the following notation: task conditions Task (perceptual, Task = 1; value-based, Task = 0); stimulation conditions TMS (pre-TMS = 0 and post-TMS = 1); V is our variable of interest, which may be behavioural or neural; and, the causal treatment effect, ϕ(V|Task, TMS), takes the following form: where E(V|Task, TMS) is the expected value of the variable of interest, V, given task and TMS condition. The first difference on the right-hand side captures the average stimulation effect for PDM while the second difference captures the average stimulation effect for VDM. The overall difference assumes that if behaviour will be the same after stimulation, then there is no effect, ϕ = 0 (Angrist and Pischke, 2009; Bertrand et al., 2004). But if there is a stimulation effect and it impairs behaviour or neural activity, then ϕ < 0.

Behavioural analyses for choice

We analysed the influence of continuous theta-burst stimulation on choice using a logit regression on choices, ρ (correct = 1, incorrect = 0) over various regressors of interest, such as TMS condition, TMS (pre-cTBS = 0, post-cTBS = 1); task, Task (perceptual = 1, value-based = 0); its interaction (Task × TMS), which measures the causal stimulation effect, ϕ; and, other regressors, Xk that we use as controls. This includes task-relevant evidence (SD for perceptual and VD for value-based, 1 to 4), response times (RTs), and task-irrelevant evidence (i.e. VD for perceptual and SD for value-based, 1 to 4). The full model is, where t indexes task, c for TMS, s for subject, and i for trial. Since our model contains a DID interaction term, nonlinearity of the logit regression results is a non-zero estimate even if the true causal effect is zero, ϕ = 0.

To remove nonlinearity bias and isolate the true causal effect, we ran another logit regression without the interaction term, and we take the difference between the two logits (Ai and Norton, 2003; Karaca-Mandic et al., 2012; Puhani, 2012), We also ran variations of the model to test for robustness, particularly GLMs with or without control variables, and we also tested robustness using various stimulation runs (see Supplementary Table 5). We used cluster-robust standard errors at the subject level under the assumption that each individual performance is independent across participants. We implemented this analysis using STATA/SE 13.1.

Behavioural analyses for response times

We similarly used DID regressions to analyse the influence of cTBS on response times (rt). Here, we simply ran a general linear model (GLM) for our regression, and we also ran variations of the model (see also Supplementary Table 5). We similarly used cluster-robust standard errors at the subject level.

Added in the revised version: Hierarchical Bayesian DDM

Earlier version: Hierarchical Bayesian DDM. We analysed the effect of cTBS on PDM and VDM using hierarchical drift diffusion model (HDDM). The model assumes a evidence is accumulated through a one-dimensional Wiener process, whereby the state of evidence, xt at time t evolves through a stochastic differential equation, Here, δ is the amount of evidence being accumulated at time t, where E represents the evidence level and k is the drift-rate that linearly scales the evidence and this is typically interpreted as quality of information processing. The model assumes evidence is accumulated at the starting point, β, and the accumulation process continues until a choice, ρ, is made at time td at a given threshold, α. Once the accumulation process terminates, the state of evidence is either xt > α (a correct decision) or xt ≤ 0 (an incorrect decision). We also accounted for visual sensory processing and motor response delays with the non-decision time parameter (nDT), τ.

Revised version: We analysed the effect of cTBS on PDM and VDM using hierarchical drift diffusion model (HDDM). The model assumes a evidence is accumulated through a one-dimensional Wiener process, whereby the state of evidence, xt at time t evolves through a stochastic differential equation, Here, δ is the amount of evidence being accumulated at time t, where E represents the evidence level and κ is the drift-rate that linearly scales the evidence and this is typically interpreted as quality of information processing. Thus, κc,s is the only free drift parameter per subject and condition; any variation in δ across evidence levels is determined by the corresponding evidence values Ec,s,i rather than by separate drift parameters for each level. When we report drift values by evidence level (Supplementary Tables 8-9), these entries summarise the posterior mean of δc,s,i for trials in each evidence bin implied by this κ × E formulation. Furthermore, the model assumes evidence is accumulated at the starting point, β, and the accumulation process continues until a choice, ρ, is made at time td at a given threshold, α. Once the accumulation process terminates, the state of evidence is either xt > α (a correct decision) or xt ≤ 0 (an incorrect decision). We also accounted for visual sensory processing and motor response delays with the non-decision time parameter (nDT), τ.

Earlier version: The hierarchical Bayesian model is implemented whereby each observed choice, yc,s,i(ρ, rt), follows a Wiener distribution, ω, where c indexes task (c = p for perceptual, c = s for value-based), s for participants (s = 1, … , Nsubjects), and i for trials (i = 1, … , Ntrials). Furthermore, the hierarchical structure contains three random variations at the trial, subject, and condition levels. We treated all interindividual differences per stimulation condition level as random effects: where N(μ, σ) is a normal distribution with mean, σ and standard deviation, σ. Here, E represents the trial-by-trial evidence levels, which we measure in absolute terms; and we fix the starting point, βc,s,i = 0.5. We used Bayesian hypothesis testing to compare posterior probability densities.

Revised version: The hierarchical Bayesian model is implemented whereby each observed choice, yc,s,i(ρ, rt), follows a Wiener distribution, τ, where c indexes task (c = p for perceptual, c = v for value-based), s for participants (s = 1, …, Nsubjects), and i for trials (i = 1, …, Ntrials). Furthermore, the hierarchical structure contains three random variations at the trial, subject, and condition levels. We treated all interindividual differences per stimulation condition level as random effects: where N(μ, σ) is a normal distribution with mean, μ and standard deviation, σ. Here, E represents the trial-by-trial evidence levels, which we measure in absolute terms; and we fix the starting point, βc,s,i = 0.5. We used Bayesian hypothesis testing to compare posterior probability densities.

Measure of accumulated evidence

We computed estimates for decision times (td(c,s)) and accumulated evidence (aEc,s) to test whether aE is a plausible representation of the accumulation process at the neural level. Following the literature (Bogacz et al., 2006, 2010), we define mean decision time as the ratio between the decision threshold and the drift rate shaped by a hyperbolic tangent function, It is important to note that reaction time, rt, is the sum of both decision and non-decision times, rt = td + τ.

We define accumulated evidence (aE) as the area below the drift process up until the accumulator reaches the decision boundary: Here, we derive aE using the area equation of a triangle, where decision time td(c,s) is the base and the decision boundary, αc,s, is the height.

MCMC sampling

To estimate all parameters, we performed Gibbs sampling via Markov Chain Montecarlo (MCMC) in JAGS (Plummer, 2016) to generate parameter posterior inferences. We drew a total of 100,000 samples from an initial burn-in step and subsequently drew a total of new 100,000 samples with three chains each. We derived each chain based on different random number generator engines with different seeds. We applied a thinning of 100 to this final sample, resulting in a final set of 1,000 samples for each parameter. This thinning assured auto-decorrelation for all latent variables of interest. We conducted Gelman-Rubin tests for each parameter to confirm chain convergence. All latent parameters in our Bayesian model had R̂ < 1.05, suggesting that all three chains converged to a target posterior distribution. We compared the difference in posterior population distributions estimated for each parameter between the stimulation conditions as well as the differences-in-differences (DID), which included differences between tasks. We tested whether the resulting distribution (i.e., the causal stimulation effect) is significantly different from zero (i.e., the null hypothesis) using the cumulative function up to or from 0 depending on the direction of the effect. We refer to this probability as Bayesian “p-values,” pmcmc.

fMRI data analysis

Participants performed eight choice-task sessions while BOLD images were recorded with a Philips Achieva 3T whole-body scanner. We used statistical parametric mapping (SPM8, Wellcome Trust Center for Neuroimaging) for image pre-processing and analysis. In particular, images were slice-time corrected (to the acquisition time of the middle slice) and realigned to account for subject’s head motion. Each participant’s T1-weighted structural image was co-registered with the mean functional image and normalised to the standard T1 MNI template using the new-segment procedure in SPM8. The functional images were normalised to the standard MNI template using the same transformation, spatially resampled to 3mm isotropic voxels, and smoothed using a Gaussian kernel (FWHM, 8mm).

We estimated two general linear models (GLMs), constructed by convolving a series of appropriately placed indicator functions with the default model of the BOLD response embedded in SPM8. GLM1 contained only two indicator functions for the onsets of PDM or VDM trials. On the other hand, GLM2 contained four indicator functions for the onsets of task (PDM and VDM trials) and stimulation (pre- or post-TMS) runs, coupled with one regressor each for parametric modulation of the BOLD response by the trialwise accumulated evidence (aE). We earlier demonstrated that the theoretical average accumulated evidence is derived from population-level as well as subject-level latent DDM parameters, by dividing the estimated decision boundary by the estimated drift rate. To construct a trialwise measure of aE, we exploit the fact that the length of the RTs is directly proportional to the size of the decision boundary while the evidence level, E, is directly proportional to the drift rate (Basten et al., 2010; Domenech et al., 2017; Kiani et al., 2014). With this mapping, we can then construct a parametric trialwise measure of accumulated evidence, aEt,c,s,i, where the square root function accounts for the concave nonlinearity in accumulated evidence. Previous work (Tajima et al., 2016) has shown theoretically that the shape of the accumulated evidence is indeed concave, where it suggests that the rate of accumulating evidence is decreasing as the decision process continues to accumulate. This concavity in aE is consistent with DDM predictions where evidence accumulation is steeper during earlier responses and begins to plateau at later responses (Ratcliff and McKoon, 2008; Ratcliff and Smith, 2004).

We convolved our GLMs with a canonical haemodynamic response function (HRF), modelled MR image autocorrelations with first-order autoregressive model, and included 6 motion parameters (obtained during realignment) as regressors of no interest. After fitting the model to the BOLD data, we tested regressors for statistical significance at the second-level, in random-effects group one-sample t-tests of the corresponding single-subject contrast images. We performed statistical inference at the cluster level, using whole-brain family-wise-error-corrected (FWE-corrected) statistical threshold of p < 0.05, based on a cluster-forming voxel cutoff at p < 0.005 (or T(19) = 2.9). For hypothesis-guided region-of-interest (ROI) analysis (i.e., left SFS stimulation site, MNI coordinates: x = −24, y = 24, z = 36), we corrected for multiple comparisons using small-volume correction (SVC, p < 0.05) restricted within a 10 mm sphere around the target coordinates. We extracted neural betas from this spherical SFS ROI for each participant to perform hypothesis testing and correlational analysis.

Functional connectivity

We ran a psychophysiological interaction (PPI) analysis (Friston et al., 1997) to investigate the changes in functional connectivity between the left SFS and other brain regions due to cTBS. Here, we extracted physiological time series in the SFS seed region, which corresponds to the time-course of the first eigenvariate across all voxels in the region using principal component analysis (Friston et al., 1993). The psychological regressor corresponded to the difference in accumulated evidence, aE (as described in GLM2) between PDM and VDM. We generated PPI estimates from the psychological regressors and the time series from the left SFS, and we then computed the PPI contrasts-of-interest for PDM and VDM. Statistical inference on subject-specific PPI maps was performed using second-level random-effects analysis across participants to allow for group-level inferences. For each participant, we also extracted PPI neural betas, which measures the degree of functional coupling between the left SFS, and we then performed hypothesis testing and correlational analysis.

Hierarchical Bayesian neural-DDM

We also analysed whether the inclusion of raw trial-by-trial BOLD response extracted from left SFS and attach it to any of the DDM parameters can improve model evidence. Such a result would suggest that neural activity in the left SFS directly related to the model’s latent decision-relevant parameters. We used z-scored single-trial neural beta estimates extracted from the left SFS target site. We implemented four a-priori models regarding the role of the left SFS on the decision parameters: Model 1 assumes that the left SFS modulated the decision threshold (Supplementary Fig. 8a), while Model 2 assumes that left SFS modulated the drift rate (Supplementary Fig. 8b): where γ is the scale parameter for trial-by-trial left SFS activity, θ. On the other hand, Models 3 and 4 assume that the left SFS modulates both boundary and drift: Model 3 assumes separate scale parameters for each latent process (see Supplementary Fig. 8c) while Model 4 assumes a common scale parameter for both boundary and drift (see Supplementary Fig. 8d). Model comparison used the Deviance Information Criterion (DIC = D̄ + pD), where pD is the effective number of parameters; thus DIC penalizes model complexity, and lower DIC denotes better predictive accuracy after accounting for complexity. We then used the best model to re-estimate our latent parameters and to perform Bayesian post-hoc inferences.

Correlating causal changes between neural, latent, and behavioural variables

We tested whether there were any correlational changes between neural, v, and behavioural, π, measures after stimulation. The marginal effect, r, measures the correlational change in neural measure, v, given behavioural measure, π, We test the marginal effect, r, of the correlational change between our neural and behavioural measures using our DID regression at trial and subject levels. With trialwise data, we used logit regression to test whether the marginal effect of trialwise changes in left SFS, Neur, will affect choices, ρ. Similar to previous models, we included various regressors of interest, especially the triple interaction (Neur × Task × TMS), which accounts for the causal TMS effect, ϕ, as well as other regressors, Xk. The full model is, To remove nonlinearity bias and isolate the true causal effect, we ran another logit regression without ϕ, and then we took the difference between the two logit models, We similarly ran a DID-GLM to test whether the marginal effect of trialwise left SFS neural betas will causally affect RTs, We also used cluster-robust standard errors at the subject level in all of our analysis.

With subject-level data, we similarly used linear mixed-effects regression models to test whether the marginal effect of subject-level neural betas (left SFS or PPI) v, will affect behavioural outcomes or DDM-latent parameters, π. Similarly, we estimated the marginal effect, ϕ, with a three-way interaction, (v × Task × TMS), This three-way interaction measures whether the correlations between neural activity (left SFS, PPI betas) and behaviour (choice, DDM parameters) are causally affected by stimulation, TMS, and whether the effect is specific only during the perceptual task.